```python
# Initialize Otter
import otter
grader = otter.Notebook("lab6.ipynb")
```

# Lab 6: Fitting Models to Data

In this lab, you will practice using a numerical optimization package `cvxpy` to compute solutions to optimization problems. The example we will use is a linear fit and a quadratic fit.

```python
import pandas as pd
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
```

## Objectives for Lab 6:

Models and fitting models to data is a common task in data science. In this lab, you will practice fitting models to data. The models you will fit are:

* Linear fit
* Normal distribution

## Boston Housing Dataset

__Warning__: The Boston housing prices dataset has an ethical problem: as investigated in [[1]](https://medium.com/%2540docintangible/racist-data-destruction-113e3eff54a8), the authors of this dataset engineered a non-invertible variable “B” assuming that racial self-segregation had a positive impact on house prices [[2]](https://www.researchgate.net/publication/4974606_Hedonic_housing_prices_and_the_demand_for_clean_air). Furthermore the goal of the research that led to the creation of this dataset was to study the impact of air quality but it did not give adequate demonstration of the validity of this assumption.
The scikit-learn maintainers therefore strongly discourage the use of this dataset unless the purpose of the code is to study and educate about ethical issues in data science and machine learning.

The warning message is from [sklearn.dataset](https://scikit-learn.org/1.0/modules/generated/sklearn.datasets.load_boston.html) documentation.

We will remove variable 'B' directly and the following analysis will not be revelant to 'B'.

```python
data_url = "http://lib.stat.cmu.edu/datasets/boston"
boston_dataset = pd.read_csv( filepath_or_buffer=data_url, delim_whitespace=True, skiprows=21, header=None)

columns = ['CRIM','ZN','INDUS','CHAS','NOX','RM','AGE','DIS','RAD', 'TAX','PTRATIO', 'B', 'LSTAT', 'MEDV']
#Flatten all the values into a single long list and remove the nulls
values_w_nulls = boston_dataset.values.flatten()
all_values = values_w_nulls[~np.isnan(values_w_nulls)]
```

```python
#Reshape the values to have 14 columns and make a new df out of them
housing = pd.DataFrame(
    data = all_values.reshape(-1, len(columns)),
    columns = columns,
)
housing.drop(columns=['B'], inplace = True)
housing.head()
```

```python
fig, ax = plt.subplots(figsize=(10, 7))
sns.scatterplot(x='LSTAT', y='MEDV', data=housing)
plt.show()
```

The model for the relationship between the response variable MEDV ($y$) and predictor variables LSTAT ($u$) and RM ($v$) is that
$$ y_i = \beta_0 + \beta_1 u_i + \epsilon_i, $$
where $\epsilon_i$ is random noise.

In order to fit the linear model to data, we minimize the sum of squared errors of all observations, $i=1,2,\dots,n$. 
$$\begin{aligned}
&\min_{\beta} \sum_{i=1}^n (y_i - \beta_0 + \beta_1 u_i )^2 = \min_{\beta} \sum_{i=1}^n (y_i - x_i^T \beta)^2 = \min_{\beta} \|y - X \beta\|_2^2
\end{aligned}$$
where $\beta = (\beta_0,\beta_1)^T$, and $x_i^T = (1, u_i)$. Therefore, $y = (y_1, y_2, \dots, y_n)^T$ and $i$-th row of $X$ is $x_i^T$. 

## Question 1: Constructing Data Variables

Define $y$ and $X$ from `housing` data.

```python
y = ...
X1 = ...
...
# X.insert(..., 'intercept', ...)
```

```python
grader.check("q1")
```

## Installing CVXPY

First, install `cvxpy` package by running the following bash command:

```python
#!pip install cvxpy
```

## Question 2: Fitting Linear Model to Data

Read this example of how cvxpy problem is setup and solved: https://www.cvxpy.org/examples/basic/least_squares.html

The usage of cvxpy parallels our conceptual understanding of components in an optimization problem:
* `beta` are the variables $\beta$
* `loss` is sum of squared errors
* `prob` minimizes the loss by choosing $\beta$

Make sure to extract the data array of data frames (or series) by using `values`: e.g., `X.values` 

```python
import cvxpy as cp

beta2 = ...
loss2 = ...
prob2 = ...

prob2.solve()

yhat2 = ...
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

## Question 3: Visualizing resulting Linear Fit

Visualize fitted model by plotting `LSTAT` by `MEDV`.

```python
fig, ax = plt.subplots(figsize=(10, 7))

...

plt.show()
```

<!-- END QUESTION -->

## Question 4: Fitting Quadratic Model to Data

Add a column of squared `LSTAT` values to `X`. The new model is,

Then, fit a quadratic model to data.

```python
X2 = X1.copy()
X2.insert(2, 'LSTAT^2', X2['LSTAT']**2)

beta4 = ...
loss4 = ...
prob4 = ...

prob4.solve()

yhat4 = ...
```

```python
grader.check("q4a")
```

<!-- BEGIN QUESTION -->

Visualize quadratic fit:

```python
fig, ax = plt.subplots(figsize=(10, 7))

...

plt.show()
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 3)
```
>>> X1.columns
Index(['intercept', 'LSTAT'], dtype='object')
```
```
>>> X1.shape
(506, 2)
```
```
>>> np.all(y == housing['MEDV'])
True
```
```
>>> np.all(X1['intercept'] == 1)
True
```
```
>>> np.all(X1['LSTAT'] == housing['LSTAT'])
True
```


Tests for q2 (points: 3)
```
>>> np.round(beta2.value, 3)
array([34.554, -0.95 ])
```


Tests for q4a (points: 3)
```
>>> np.round(beta4.value, 3)
array([42.862, -2.333,  0.044])
```
