```python
# Initialize Otter
import otter
grader = otter.Notebook("lab7.ipynb")
```

# Lab 7: Regular Expressions 

```python
import pandas as pd
import numpy as np
import re

from pathlib import Path
```

## Objectives:

First, read this DS 100 text book section on regular expressions: <https://www.textbook.ds100.org/ch/08/text_regex.html>
 
Then, you will practice the basic usage of regular expressions and also learn to use `re` module in Python. Some of the materials are based on the tutorial at <http://opim.wharton.upenn.edu/~sok/idtresources/python/regex.pdf>. As you work through the first part of the lab, you may also find the website <http://regex101.com> helpful.

You will likely run out of time; however, you can finish the exercises on your own.

---
# Part 1: Regular Expressions

We'll start by learning about the simplest possible regular expressions. Since regular expressions are used
to operate on strings, we'll start with the most common task: matching characters.

Most letters and characters will simply match themselves. For example, the regular expression `r'test'` will match the string `test` exactly. There are exceptions to this rule; some characters are special, and don't match themselves.

Here is a list of metacharacters that are widely used in regular experssion. 

|Pattern|Description|
|--- |--- |
|`^`|Matches beginning of line.|
|`$`|Matches end of line.|
|`.`|Matches any single character except newline.|
|`*`|Matches 0 or more occurrences of preceding expression.|
|`+`|Matches 1 or more occurrence of preceding expression.|
|`?`|Matches 0 or 1 occurrence of preceding expression.|
|`[...]`|Matches any single character in brackets.|
|`[^...]`|Matches any single character not in brackets.|
|`{n}`|Matches exactly n number of occurrences of preceding expression.|
|`{n,}`|Matches n or more occurrences of preceding expression.|
|`{n,m}`|Matches at least n and at most m occurrences of preceding expression.|
|`a\|b`|Matches either a or b.|
|`\1...\9`|Matches n-th grouped subexpression.|

Perhaps the most important metacharacter is the backslash, ‘\’. As in Python string literals, the backslash
can be followed by various characters to signal various special sequences. It’s also used to escape all the
metacharacters so you can still match them in patterns; for example, if you need to match a `[` or `\`, you
can precede them with a backslash to remove their special meaning:  `\[` or `\\`. 

The following predefined special sequences are available:

|Pattern|Description|
|--- |--- |
|`\d`|Matches any decimal digit; this is equivalent to the class `[0-9]`|
|`\D`|Matches any non-digit character; this is equivalent to the class `[^0-9]`.|
|`\s`|Matches any whitespace character; this is equivalent to the class `[ \t\n\r\f\v]`|
|`\S`|Matches any non-whitespace character; this is equivalent to the class `[^ \t\n\r\f\v]`.|
|`\w`|Matches any alphanumeric character; this is equivalent to the class `[a-zA-Z0-9_]`|
|`\W`|Matches any non-alphanumeric character; this is equivalent to the class `[^a-zA-Z0-9_]`.|

# Question 1
In this question, write patterns that match the given sequences. It may be as simple as the common letters on each line.


## Question 1a

Write a single regular expression to match the following strings without using the `|` operator.

1. **Match:** `abcdefg`
1. **Match:** `abcde`
1. **Match:** `abc`
1. **Skip:** `c abc`

<!-- grade_id: q1a-answer, points: 0, grade: false, solution: true -->
```python
regx1 = r"" # fill in your pattern
```

```python
grader.check("q1a")
```


## Question 1b

Write a single regular expression to match the following strings without using the `|` operator.

1. **Match:** `can`
1. **Match:** `man`
1. **Match:** `fan`
1. **Skip:** `dan`
1. **Skip:** `ran`
1. **Skip:** `pan`


<!-- grade_id: q1b-answer, points: 0, grade: false, solution: true -->
```python
regx2 = r"" # fill in your pattern
```

```python
grader.check("q1b")
```

# Question 2

Now that we have written a few regular expressions, we are now ready to move beyond matching. In this question, we'll take a look at some methods from the `re` package.


## Question 2a:

Write a Python program to extract and print the numbers of a given string. 

1. **Hint:** use `re.findall`
2. **Hint:** use `\d` for digits and one of either `*` or `+`.


<!-- grade_id: q2a-answer, points: 0, grade: false, solution: true -->
```python
text_q2a = "Ten 10, Twenty 20, Thirty 30"

res_q2a = ...

res_q2a
```

```python
grader.check("q2a")
```


## Question 2b:

Write a Python program to replace at most 2 occurrences of space, comma, or dot with a colon.

**Hint:** use `re.sub(regex, "newtext", string, number_of_occurences)`

<!-- grade_id: q2b-answer, points: 0, grade: false, solution: true -->
```python
text_q2b = 'Python Exercises, PHP exercises.'
res_q2b = ... # Hint: use re.sub()

res_q2b
```

```python
grader.check("q2b")
```


## Question 2c: 

Write a Python program to extract values between quotation marks of a string.

<!-- grade_id: q2c-answer, points: 0, grade: false, solution: true -->
```python
text_q2c = '"Python", "PHP", "Java"'
res_q2c = ... # Hint: use re.findall()

res_q2c
```

```python
grader.check("q2c")
```

## Question 2d:

Write a regular expression to extract and print the quantity and type of objects in a string. You may assume that a space separates quantity and type, ie. `"{quantity} {type}"`. See the example string below for more detail.

1. **Hint:** use `re.findall`
2. **Hint:** use `\d` for digits and one of either `*` or `+`.

```python
text_q2d = "I've got 10 eggs that I stole from 20 gooses belonging to 30 giants in 40 days."

res_q2d = ...

res_q2d
```

```python
grader.check("q2d")
```

<!-- BEGIN QUESTION -->

## Question 2e (optional):

Write a regular expression to replace all words that are not `"mushroom"` with `"badger"`.


```python

text_qe = 'this is a word mushroom mushroom'
res_qe = ... # Hint: https://www.regextester.com/94017
res_qe
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Download the zip file and submit to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1a (points: 4)
```
>>> '|' not in regx1
True
```
```
>>> re.search(regx1, 'abc').group()
'abc'
```
```
>>> re.search(regx1, 'abcde').group()
'abcde'
```
```
>>> re.search(regx1, 'abcdefg').group()
'abcdefg'
```
```
>>> re.search(regx1, 'c abc') is None
True
```


Tests for q1b (points: 4)
```
>>> '|' not in regx2
True
```
```
>>> re.match(regx2, 'can').group()
'can'
```
```
>>> re.match(regx2, 'fan').group()
'fan'
```
```
>>> re.match(regx2, 'man').group()
'man'
```
```
>>> re.match(regx2, 'dan') is None
True
```
```
>>> re.match(regx2, 'ran') is None
True
```
```
>>> re.match(regx2, 'pan') is None
True
```


Tests for q2a (points: 4)
```
>>> res_q2a
['10', '20', '30']
```


Tests for q2b (points: 4)
```
>>> res_q2b
'Python:Exercises: PHP exercises.'
```


Tests for q2c (points: 4)
```
>>> res_q2c
['Python', 'PHP', 'Java']
```


Tests for q2d (points: 4)
```
>>> assert res_q2d == ['10 eggs', '20 gooses', '30 giants', '40 days']
```
