```python
# Initialize Otter
import otter
grader = otter.Notebook("lab3.ipynb")
```

# Lab 3: Data Cleaning and Visualization

In this lab, you will be working with a dataset from the City of Berkeley containing data on calls to the Berkeley Police Department. Information about the dataset can be found [at this link](https://data.cityofberkeley.info/Public-Safety/Berkeley-PD-Calls-for-Service/k2nh-s5h5).

## Setup

Note that after activating matplotlib to display figures inline via the IPython magic `%matplotlib inline`, we configure a custom default figure size. Virtually every default aspect of matplotlib [can be customized](https://matplotlib.org/users/customizing.html).

```python
import pandas as pd
import numpy as np
import zipfile
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns

%matplotlib inline
```

```python
plt.rcParams['figure.figsize'] = (12, 9)
```

# Part 1: Cleaning and Exploring the Data

Unzipping the file named `lab3_data.zip` creates a new directory named `data` where the data will reside.

```python
!ls -alh . ## what does -alh do? 
```

 Unzip the file using `unzip` shell command.

```python
!unzip -o lab3_data.zip -d data ## what does -d do?
```

You can also use Python to unzip the archived file.

```python
import zipfile
with zipfile.ZipFile('lab3_data.zip', 'r') as zip_ref:
    zip_ref.extractall('data')
```

Let's now load the CSV file we have into a `pandas.DataFrame` object.

```python
calls = pd.read_csv("data/Berkeley_PD_-_Calls_for_Service.csv")
calls.head()
```

We see that the fields include a case number, the offense type, the date and time of the offense, the "CVLEGEND" which appears to be related to the offense type, a "CVDOW" which has no apparent meaning, a date added to the database, and the location spread across four fields.

Let's also check some basic information about these files using the `DataFrame.describe` and `DataFrame.info` methods.

```python
calls.info()
calls.describe()
```

Notice that the functions above reveal type information for the columns, as well as some basic statistics about the numerical columns found in the DataFrame. However, we still need more information about what each column represents. Let's explore the data further in Question 1.

Before we go over the fields to see their meanings, the cell below will verify that all the events happened in Berkeley by grouping on the `City` and `State` columns. You should see that all of our data falls into one group.

```python
calls.groupby(["City","State"]).count()
```

## Question 1
Above, when we called `head`, it seemed like `OFFENSE` and `CVLEGEND` both contained information about the type of event reported. What is the difference in meaning between the two columns? One way to probe this is to look at the `value_counts` for each Series.

```python
calls['OFFENSE'].value_counts().head(10)
```

```python
calls['CVLEGEND'].value_counts().head(10)
```

### Question 1a

Above, it seems like `OFFENSE` is more specific than `CVLEGEND`, e.g. "LARCENY" vs. "THEFT FELONY (OVER $950)". For those of you who don't know the word "larceny", it's a legal term for theft of personal property.

To get a sense of how many subcategories there are for each `OFFENSE`, set `calls_by_cvlegend_and_offense` equal to a multi-indexed series where the data is first indexed on the `CVLEGEND` and then on the `OFFENSE`, and the data is equal to the number of offenses in the database that match the respective `CVLEGEND` and `OFFENSE`. For example, calls_by_cvlegend_and_offense["LARCENY", "THEFT FROM PERSON"] should return 24.

Use methods `.groupby()` and `.size()`. Refer to [Hierarchical Indexes example](https://pandas.pydata.org/pandas-docs/version/0.25.3/reference/api/pandas.DataFrame.groupby.html#pandas.DataFrame.groupby) for how to group using multiple columns.


<!-- grade_id: q1a-answer, points: 0, grade: false, solution: true -->
```python
calls_by_cvlegend_and_offense = ...


```

```python
grader.check("q1a")
```

### Question 1b

In the cell below, set `answer1b` equal to a list of strings corresponding to the possible values for `OFFENSE` when `CVLEGEND` is "LARCENY". You can type the answer manually, or you can create an expression that automatically extracts the names.




<!-- grade_id: q1b-answer, points: 0, grade: false, solution: true -->
```python
# You may use this cell for your scratch work as long as you enter
# in your final answers in the answer1b variable.
answer1b = ... 


```

```python
grader.check("q1b")
```

## Question 2

What are the five crime types of CVLEGEND that have the most crime events? You may need to use `value_counts` to find the answer.
Save your results into `answer2` as a list of strings.

**Hint:** *The `keys` method of the Series class might be useful.*



<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
answer2 = ...




```

```python
grader.check("q2")
```

---
# Part 2: Visualizing the Data

## Pandas vs. Seaborn Plotting

Pandas offers basic functionality for plotting. For example, the `DataFrame` and `Series` classes both have a `plot` method. However, the basic plots generated by pandas are not particularly pretty. While it's possible to manually use matplotlib commands to make pandas plots look better, we'll instead use a high level plotting library called Seaborn that will take care of most of this for us.

As you learn to do data visualization, you may find the [pandas documentation](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.plot.html) and [Seaborn documentation](https://seaborn.pydata.org/api.html) helpful!

As an example of the built-in plotting functionality of pandas, the following example uses `plot` method of the `Series` class to generate a `barh` plot type to visually display the value counts for `CVLEGEND`.

```python
ax = calls['CVLEGEND'].value_counts().plot(kind='barh')
ax.set_ylabel("Crime Category")
ax.set_xlabel("Number of Calls")
ax.set_title("Number of Calls By Crime Type");
```

By contrast, the Seaborn library provides a specific function `countplot` built for plotting counts. It operates directly on the DataFrame itself i.e. there's no need to call `value_counts()` at all. This higher level approach makes it easier to work with. Run the cell below, and you'll see that the plot is much prettier (albeit in a weird order).

```python
ax = sns.countplot(data=calls, y="CVLEGEND")
ax.set_ylabel("Crime Category")
ax.set_xlabel("Number of Calls")
ax.set_title("Number of Calls By Crime Type");
```

If we want the same ordering that we had in the pandas plot, we can use the order parameter of the `countplot` method. It takes a list of strings corresponding to the axis to be ordered. By passing the index of the `value_counts`, we get the order we want.

```python
ax = sns.countplot(data=calls, y="CVLEGEND", order=calls["CVLEGEND"].value_counts(ascending=True).index);
ax.set_ylabel("Crime Category")
ax.set_xlabel("Number of Calls")
ax.set_title("Number of Calls By Crime Type");
```

Voilà! Now we have a pretty bar plot with the bars ordered by size. Though `seaborn` appears to provide a superior plot from a aesthetic point of view, the `pandas` plotting library is also good to understand. You'll get practice using both libraries in the following questions.

## An Additional Note on Plotting in Jupyter Notebooks

You may have noticed that many of our code cells involving plotting end with a semicolon (;). This prevents any extra output from the last line of the cell that we may not want to see. Try adding this to your own code in the following questions!

## Question 3

Now it is your turn to make some plots using `pandas` and `seaborn`. Let's start by looking at the distribution of calls over days of the week.

The CVDOW field isn't named helpfully and it is hard to see the meaning from the data alone. According to the website linked at the top of this notebook, CVDOW is actually indicating the day that events happened. 0->Sunday, 1->Monday ... 6->Saturday. 


### Question 3a

Add a new column `Day` into the `calls` dataframe that has the string weekday (eg. 'Sunday') for the corresponding value in CVDOW. For example, if the first 3 values of `CVDOW` are `[3, 6, 0]`, then the first 3 values of the `Day` column should be `["Wednesday", "Saturday", "Sunday"]`.

**Hint:** *Try using the [Series.map](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.Series.map.html) function on `calls["CVDOW"]`.  Can you assign this to the new column `calls["Day"]`?*


```python
days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"]
day_indices = range(7)
indices_to_days_dict = dict(zip(day_indices, days)) # Should look like {0:"Sunday", 1:"Monday", ..., 6:"Saturday"}
```

<!-- grade_id: q3a-answer, points: 0, grade: false, solution: true -->
```python
...



```

```python
grader.check("q3a")
```

### Question 3b

Run the cell below to create a `seaborn` plot. This plot shows the number of calls for each day of the week. Notice the use of the `rotation` argument in `ax.set_xticklabels`, which rotates the labels by 90 degrees.

```python
ax = sns.countplot(data=calls, x='Day', order=days)
ax.set_xticklabels(ax.get_xticklabels(), rotation=90)
ax.set_ylabel("Number of Calls")
ax.set_title("Number of Calls For Each Day of the Week");
```

Now, let's make the same plot using `pandas`. Construct a vertical bar plot with the count of the number of calls (entries in the table) for each day of the week **ordered by the day of the week** (eg. `Sunday`, `Monday`, ...). Do not use `sns` for this plot. Be sure that your axes are labeled and that your plot is titled.

**Hint:** *Given a series `s`, and an array `coolIndex` that has the same entries as in `s.index`, `s[coolIndex]` will return a copy of the series in the same order as `coolIndex`.*



<!-- grade_id: q3b-answer, points: 0, grade: false, solution: true -->
```python

...
# LEAVE BELOW CODE/NEEDED FOR SOLUTION CHECKING!
ax_3b = plt.gca();
```

```python
grader.check("q3b")
```

## Question 4

It seems weekdays generally have slightly more calls than Saturday or Sunday, but the difference does not look significant.  

We can break down into some particular types of events to see their distribution. For example, let's make a bar plot for the CVLEGEND "NOISE VIOLATION". Which day is the peak for "NOISE VIOLATION"?


### Question 4a

This time, use `seaborn` to create a vertical bar plot of the number of total noise violations reported on each day of the week, again ordered by the days of the week starting with Sunday. Do not use `pandas` to plot.

**Hint:** *If you're stuck, use the code for the seaborn plot in Question 3b as a starting point.*


<!-- grade_id: q4a-answer, points: 0, grade: false, solution: true -->
```python


...
# LEAVE BELOW CODE/NEEDED FOR SOLUTION CHECKING
ax_4a = plt.gca();

```

```python
grader.check("q4a")
```

<!-- BEGIN QUESTION -->

### Question 4b

Do you realize anything interesting about the distribution of NOISE VIOLATION calls over a week? Type a 1-sentence answer in the cell below.


<!-- grade_id: q4b-answer, points: 1, grade: true, solution: true -->
**SOLUTION**

<!-- END QUESTION -->

## Question 5

Let's look at a similar distribution but for a crime we have much more calls data about. In the cell below, create the same plot as you did in Question 4, but now looking at instances of the CVLEGEND "FRAUD" (instead of "NOISE VIOLATION"). Use either `pandas` or `seaborn` plotting as you desire.


<!-- grade_id: q5-answer, points: 0, grade: false, solution: true -->
```python

...
# Leave this for grading purposes
ax_5 = plt.gca()
```

```python
grader.check("q5")
```

## Question 6

### Question 6a

Now let's look at the EVENTTM column which indicates the time for events. Since it contains hour and minute information, let's extract the hour info and create a new column named `Hour` in the `calls` dataframe. You should save the hour as an `int`. Then plot the frequency of each hour in the table (i.e., `value_counts()`) sorted by the hour of the day (i.e., `sort_index()`).

You will want to look into how to use:

* [Series.str.slice](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.Series.str.slice.html#pandas.Series.str.slice) to select the substring.
* [Series.astype](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.Series.astype.html) to change the type.

**Hint:** *The `str` helper member of a series can be used to grab substrings.  For example, `calls["EVENTTM"].str.slice(3,5)` returns the minute of each hour of the `EVENTTM`.*


<!-- grade_id: q6a-answer, points: 0, grade: false, solution: true -->
```python
...
```

```python
grader.check("q6a")
```

The code in the cell below creates a pandas bar plot showing the number of FRAUD crimes committed at each hour of the day.

```python
ax = calls[calls["CVLEGEND"] == "FRAUD"]['Hour'].value_counts().sort_index().plot(kind='bar')
ax.set_xlabel("Hour of the Day")
ax.set_ylabel("Number of Calls")
ax.set_title("Number of Calls Reporting Fraud For Each Day of the Week");
```

The cell below contains a seaborn plot of the same data.

<!-- grade_id: q6b-answer, points: 0, grade: false, solution: true -->
```python
ax = sns.countplot(calls[calls["CVLEGEND"] == "FRAUD"],x='Hour')
ax.set_xlabel("Hour of the Day")
ax.set_ylabel("Number of Calls")
ax.tick_params(rotation=90,axis='x')
ax.set_title("Number of Calls Reporting Fraud For Each Day of the Week");

```

<!-- BEGIN QUESTION -->

### Question 6b

According to our plots, there seems to be a spike in calls reporting fraud at midnight. Do you trust that this spike is legitimate, or could there be an issue with our data? Explain your reasoning in 1-2 sentences in the cell below.


<!-- grade_id: q6c-answer, points: 1, grade: true, solution: true -->
**SOLUTION**

<!-- END QUESTION -->

_Cell Intentionally Blank_

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1a (points: 3)
```
>>> calls_by_cvlegend_and_offense['LARCENY', 'THEFT FROM PERSON'] == 13
True
```


Tests for q1b (points: 3)
```
>>> isinstance(answer1b, list)
True
```
```
>>> all([isinstance(elt, str) for elt in answer1b])
True
```
```
>>> len(answer1b) == 3
True
```
```
>>> all([elt in calls['OFFENSE'].values for elt in answer1b])
True
```
```
>>> set([a.strip().upper() for a in answer1b]) == set(['THEFT FELONY (OVER $950)', 'THEFT FROM PERSON', 'THEFT MISD. (UNDER $950)'])
True
```


Tests for q2 (points: 3)
```
>>> isinstance(answer2, list)
True
```
```
>>> all([isinstance(elt, str) for elt in answer2])
True
```
```
>>> all([elt in calls['CVLEGEND'].values for elt in answer2])
True
```
```
>>> set(answer2) == set(['LARCENY', 'BURGLARY - VEHICLE', 'VANDALISM', 'DISORDERLY CONDUCT', 'ASSAULT'])
False
```


Tests for q3a (points: 3)
```
>>> set(calls['Day']) == {'Friday', 'Monday', 'Saturday', 'Sunday', 'Thursday', 'Tuesday', 'Wednesday'}
True
```
```
>>> list(calls['Day'][:5]) == ['Wednesday', 'Wednesday', 'Friday', 'Tuesday', 'Saturday']
False
```


Tests for q3b (points: 3)
```
>>> import matplotlib
>>> np.alltrue(np.array([l.get_text() for l in ax_3b.xaxis.get_ticklabels()]) == days)
True
```
```
>>> import matplotlib
>>> bars = [rect.get_height() for rect in ax_3b.get_children() if isinstance(rect, matplotlib.patches.Rectangle) and rect.get_x() != 0.0]
>>> np.allclose(np.array(bars)[-7:], calls['Day'].value_counts()[days].values)
True
```


Tests for q4a (points: 3)
```
>>> all([l.get_text() == d for l, d in zip(ax_4a.get_xticklabels(), days)])
True
```
```
>>> bars = [rect.get_height() for rect in ax_4a.get_children() if isinstance(rect, matplotlib.patches.Rectangle) and rect.get_x() != 0.0]
>>> np.allclose(np.array(bars)[-3:], [1, 1, 2])
False
```


Tests for q5 (points: 3)
```
>>> import matplotlib
>>> np.alltrue(np.array([l.get_text() for l in ax_5.xaxis.get_ticklabels()]) == days)
True
```
```
>>> import matplotlib
>>> bars = [rect.get_height() for rect in ax_5.get_children() if isinstance(rect, matplotlib.patches.Rectangle) and rect.get_x() != 0.0]
>>> np.allclose(np.array(bars)[-7:], np.array([27, 32, 29, 33, 35, 40, 12]))
False
```


Tests for q6a (points: 3)
```
>>> 'Hour' in calls.columns
True
```
```
>>> set(calls['Hour']) == set(range(24))
True
```
```
>>> list(calls['Hour'][:5]) == [18, 16, 10, 13, 12]
True
```
```
>>> np.allclose(calls['Hour'].mean(), 13.150702)
True
```
