```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment2.ipynb")
```

# Assignment 2: Exploratory Data Analysis in Professional Basketball

In this assignment we'll conduct an exploratory data analysis of professional basketball data.  Basketball is a team sport in which the goal is to try to outscore the amount in a fixed amount of time. Points are scored (either 2 or 3 points) by putting the ball throw a hoop on one end of the court.  An attempt at putting the ball throw the hoop is known as a "shot".  If helpful, you can read more about [the rules of basketball](https://en.wikipedia.org/wiki/Rules_of_basketball).

The National Basketball Association (NBA) is the professional basketball league in the United States and provides a nice website with many statistics gathered on teams and players in the league: [http://stat.nba.com](http://stat.nba.com). 

<!-- BEGIN QUESTION -->

## Question 1: Managing data files

We will use data that is available from NBA. Although NBA doesn't officially make the data API (application programming interface) public, people have figured out ways to access their data programmatically ([1](https://github.com/swar/nba_api), [2](https://github.com/seemethere/nba_py/wiki/stats.nba.com-Endpoint-Documentation)). However, NBA does not offer an official API and it is possible to get our JupyterHub blocked by the site if we use them. Therefore, in this assignment, the raw data downloads are provided to you in a zip file: https://ucsb.box.com/shared/static/z6y3etgikbzbnf0ld4brvc95xtgjcrie.zip

Download and unzip the file to a directory named `data` using command line commands (unzipping on Windows and Mac may not work because different OS have different constraints on filename lengths, etc.). Adding an exclamation point in the Jupyter notebook cell indicates that `bash` shell interpreter will execute your command.
```bash
wget -nc https://ucsb.box.com/shared/static/z6y3etgikbzbnf0ld4brvc95xtgjcrie.zip -O nba-data.zip
unzip -o nba-data.zip -d data
```
What these commands are doing:
* `wget` downloads files ([what do each of the pieces do?](https://explainshell.com/explain?cmd=wget+-nc+https%3A%2F%2Fucsb.box.com%2Fshared%2Fstatic%2F940qiuxyp798gv4nx8iwvazu9qqjs37e.zip+-O+nba-data.zip))
* `unzip` will unzip `nba-data.zip` into a directory named `data` (specified by `-d data`) and will overwrite any same filenames when extracting (specified by `-o`).

```python
# Run your commands in this cell
```

<!-- END QUESTION -->

After unzipping the files, you will find three types of files in `data/` directory:

* Team data: `commonTeamYears?LeagueID=00&Season=2018-19`
* Player data: `commonallplayers?LeagueID=00&Season=2018-19&IsOnlyCurrentSeason=0`
* Player's shot data: `shotchartdetail?PlayerID=[PlayerID]&PlayerPosition=&Season=2018-19&ContextMeasure=FGA&DateFrom=&DateTo=&GameID=&GameSegment=&LastNGames=0&LeagueID=00&Location=&Month=0&OpponentTeamID=0&Outcome=&Period=0&Position=&RookieYear=&SeasonSegment=&`

Each player's shot data is identified by replacing `[PlayerID]` with their numeric ID.

Here is how we will read in the data:
* Each data file contains text in [JSON (Javascript Object Notation) format](https://www.w3schools.com/python/python_json.asp).
* First, read the data content as text (using [Path.read_text()](https://docs.python.org/3/library/pathlib.html#pathlib.Path.read_text) from `pathlib` module)
* Second, we convert it to a Python dictionary format (using [json.loads()](https://docs.python.org/3/library/json.html#json.loads) in `json` module)
* Third, identify DataFrame content 
* Fourth, identify DataFrame header
* Fifth, assemble DataFrame

Another way to unzip a file is using `zipfile`.

```python
import zipfile
with zipfile.ZipFile('nba-data.zip', 'r') as zip_ref:
    zip_ref.extractall('data')
```

### Question 1a: Team Data

Read team data file into a pandas data frame named `allteams` starting from the given code below.

```python
from pathlib import Path
import json
import pandas as pd
import numpy as np

fname  = 'data/commonTeamYears?LeagueID=00&Season=2018-19' # directory_name/file_name
step_1 = Path(fname).read_text()              # str
step_2 = json.loads(step_1)                   # dict
step_3 = step_2['resultSets'][0]['rowSet']    # list 
step_4 = step_2['resultSets'][0]['headers']   # list
```

```python
# print out each of step_1 through step_4 and understand what each line does
```

Use variables constructed above to assemble `allteams` DataFrame.

Drop any teams that no longer exist as of 2019. These teams show None in `ABBREVIATION` column.



```python
allteams = ...
```

```python
grader.check("q1a")
```

### Question 1b: Player Data

`pathlib` has flexible ways to specify file and directory paths. For example, the following are equivalent:

* `Path('data/commonallplayers?LeagueID=00&Season=2018-19&IsOnlyCurrentSeason=0')`
* `Path('data') / 'commonallplayers?LeagueID=00&Season=2018-19&IsOnlyCurrentSeason=0')`
* `Path('data').joinpath('commonallplayers?LeagueID=00&Season=2018-19&IsOnlyCurrentSeason=0')`

Read players data file with name `data/commonallplayers?LeagueID=00&Season=2018-19&IsOnlyCurrentSeason=0`. 
Assemble pandas DataFrame with name `allplayers`. Set row index to be `PERSON_ID` and `sort_index`.


```python
dirname  = 'data' # directory_name
filename = 'commonallplayers?LeagueID=00&Season=2018-19&IsOnlyCurrentSeason=0' # file_name
step_1 = ...
...
allplayers = ...
```

```python
grader.check("q1b")
```

### Question 1c: Shots Data

`pathlib` can also find all filenames that match a given pattern using [`Path.glob()` method](https://docs.python.org/3/library/pathlib.html#pathlib.Path.glob).

For example, teams data and players data start with the pattern `common` followed by a wildcard `*`: `common*`. 

We can use this to retrieve two file names with one call:

```python
two_files = Path('data').glob('common*') # generator: https://www.educative.io/edpresso/generator-vs-iterator-in-python
list(two_files)                          # list
```

All file names for shots data start with `shotchartdetail`. 

Use this as the pattern to 
* First, read all file names into `allshots_files`
* Second, loop over each file in `allshots_files` and assemble a dataframe
* Third, add as an element in a list named `allshots_list` (each file is an data frame item in the list). 
* Fourth, concatenate all dataframes into one dataframe named `allshots`. Set the row index to be `PLAYER_ID` and `sort_index`.


```python
allshots_files = ...
allshots_files.sort()
allshots_list  = list()

for f in allshots_files:
    ...

allshots = ...
```

```python
grader.check("q1c")
```

### Question 1d: Extract Stephen Curry's Shot Data

Use [`allplayers.query()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.query.html) to find the player id (index) associated with the player named "[Stephen Curry](https://en.wikipedia.org/wiki/Stephen_Curry)".  Set the value of `PlayerID` as `curry_id` of type `str`. 

Subset all of Stephen Curry's shots in a data frame named `curry_data`. Also, set the dtype of `SHOT_MADE_FLAG` to `'bool'` in one command. Something like:  
```
curry_data = allshots.query(???).astype(????)
```



```python
# fill-in all ...
query_str = '... == "Stephen Curry"'
curry_id = str(allplayers.query(...).index.values[...])
curry_data = allshots.query('PLAYER_ID == ' + curry_id).astype({... : ...})
...
```

```python
grader.check("q1d")
```

## Question 2: Visualization

### Question 2a: All Shots Scatter Plot
Use `seaborn` to create scatter plot of the location of Stephen Curry's shot attempts from this year (`LOC_X` and `LOC_Y`).  When you call a scatterplot, seaborn returns a figure in an object, we'll call it `ax`.  We can set properties of the figure by calling methods on `ax`.  Use this approach to set the x-axis limits to span (-300, 300), the y-axis limits to span (-100, 500).

```python
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns

plt.figure(figsize=[12, 11])
ax2a = ...

# Set x/y limits and labels
...
plt.show()
```

```python
grader.check("q2a")
```

Understanding any dataset is difficult without context.  Lets add some important context by adding the relevant court lines into our diagram.  If you are interested, you can read more about the lines and dimensions on the [NBA basketball court](https://en.wikipedia.org/wiki/Basketball_court).  We will use code from [http://savvastjortjoglou.com/nba-shot-sharts.html](http://savvastjortjoglou.com/nba-shot-sharts.html) to add the court markings to our diagram.  The `draw_court` function below will do this for us.  The below cell will generate an example court.

```python
## code is from http://savvastjortjoglou.com/nba-shot-sharts.html
def draw_court(ax=None, color='black', lw=1, outer_lines=False):
    
    from matplotlib.patches import Circle, Rectangle, Arc
    from matplotlib.pyplot import gca
    
    # If an axes object isn't provided to plot onto, just get current one
    if ax is None:
        ax = gca()

    # Create the various parts of an NBA basketball court

    # Create the basketball hoop
    # Diameter of a hoop is 18" so it has a radius of 9", which is a value
    # 7.5 in our coordinate system
    hoop = Circle((0, 0), radius=7.5, linewidth=lw, color=color, fill=False)

    # Create backboard
    backboard = Rectangle((-30, -7.5), 60, 0, linewidth=lw, color=color)

    # The paint
    # Create the outer box 0f the paint, width=16ft, height=19ft
    outer_box = Rectangle((-80, -47.5), 160, 190, linewidth=lw, color=color,
                          fill=False)
    # Create the inner box of the paint, widt=12ft, height=19ft
    inner_box = Rectangle((-60, -47.5), 120, 190, linewidth=lw, color=color,
                          fill=False)

    # Create free throw top arc
    top_free_throw = Arc((0, 142.5), 120, 120, theta1=0, theta2=180,
                         linewidth=lw, color=color, fill=False)
    # Create free throw bottom arc
    bottom_free_throw = Arc((0, 142.5), 120, 120, theta1=180, theta2=0,
                            linewidth=lw, color=color, linestyle='dashed')
    # Restricted Zone, it is an arc with 4ft radius from center of the hoop
    restricted = Arc((0, 0), 80, 80, theta1=0, theta2=180, linewidth=lw,
                     color=color)

    # Three point line
    # Create the side 3pt lines, they are 14ft long before they begin to arc
    corner_three_a = Rectangle((-219, -47.5), 0, 140, linewidth=lw,
                               color=color)
    corner_three_b = Rectangle((219, -47.5), 0, 140, linewidth=lw, color=color)
    # 3pt arc - center of arc will be the hoop, arc is 23'9" away from hoop
    # I just played around with the theta values until they lined up with the 
    # threes
    three_arc = Arc((0, 0), 475, 475, theta1=22.5, theta2=157.5, linewidth=lw,
                    color=color)

    # Center Court
    center_outer_arc = Arc((0, 422.5), 120, 120, theta1=180, theta2=0,
                           linewidth=lw, color=color)
    center_inner_arc = Arc((0, 422.5), 40, 40, theta1=180, theta2=0,
                           linewidth=lw, color=color)

    # List of the court elements to be plotted onto the axes
    court_elements = [hoop, backboard, outer_box, inner_box, top_free_throw,
                      bottom_free_throw, restricted, corner_three_a,
                      corner_three_b, three_arc, center_outer_arc,
                      center_inner_arc]

    if outer_lines:
        # Draw the half-court line, baseline and side-out bound lines
        outer_lines = Rectangle((-250, -47.5), 500, 470, linewidth=lw,
                                color=color, fill=False)
        court_elements.append(outer_lines)

    # Add the court elements onto the axes
    for element in court_elements:
        ax.add_patch(element)

    return ax

plt.figure(figsize=(12,11))
draw_court(outer_lines=True)
plt.xlim(-300,300)
plt.ylim(-100,500)
plt.show()
```

<!-- BEGIN QUESTION -->

### Question 2b: All Shots Scatter Plot + Court Outline

Again use seaborn to make a scatter plot of Stephen Curry's shots. Again, set the x-axis limits to span (-300, 300), the y-axis limits to span (-100, 500) color the points by whether the shot was made or missed.  Set the missed shots to have an 'x' symbol and made shots to be a circular symbol. Call the `draw_court` function with `outer_lines` set to to be true. Save the `Axes` returned by the plot call in a variable called `ax`.


```python
plt.figure(figsize=(12, 11))
markers = {0 : "X", 1 : "o"}
ax = ...

...
plt.show()
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2c: Analyzing the Visualization

In a few sentences, discuss what makes this an effective or ineffective visualization for understanding the types of shots that Stephen Curry likes to take and is good at taking, relative to other players in the league.  Are there ways it can be improved?

**SOLUTION**

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2d: A Hexbin plot

Visualize Stephen Curry's shots by using a [hexbin plot with marginal histograms](https://seaborn.pydata.org/examples/hexbin_marginals.html). Also refer to setting [figure aesthetics](https://seaborn.pydata.org/tutorial/aesthetics.html) for what commands below do.

```python
sns.set_style("white")
joint_shot_chart = ...
joint_shot_chart.fig.set_size_inches(12,11)

# A joint plot has 3 Axes, the first one called ax_joint 
# is the one we want to draw our court onto and adjust some other settings
ax = joint_shot_chart.ax_joint
draw_court(ax, outer_lines=True)

# Adjust the axis limits and orientation of the plot in order
# to plot half court, with the hoop by the top of the plot
ax.set_xlim(-300, 300)
ax.set_ylim(500, -100)

# Get rid of axis labels and tick marks
ax.set_xlabel('')
ax.set_ylabel('')
ax.tick_params(labelbottom=False, labelleft=False)

# Add a title
ax.set_title('Stephen Curry, 2018-19, FGA', 
             y=1.2, fontsize=10)

# Add Data Source and Author
ax.text(-250,445,'\n This plot is based on code by Savvas Tjortjoglou (savvastjortjoglou.com)',
        fontsize=12);
```

<!-- END QUESTION -->

## Question 3: Binning and Smoothing Shots

So far, in we have worked with dataframes which represent each shot as a single observation (row) within the dataset.  However, this isn't a convenient data structure for the kinds of spatial analyses we will pursue below.

In this part, we will divide the court into square regions and create a matrix which includes the number of shots taken by a player in that region.  We divide the court up into square bins (i.e. a 2d histogram) and, for each player, count number of shots that fall into each bin. Fortunately, this function is relatively simple to write using `numpy` module. 

### Question 3a: 2D Smoothing

Fill in the `bin_shots` function below. Use `np.histgram2d` to count count the shots in each bin.  The bins are defined `bin_edges` which is a pandas Series of the form `(xedges, yedges)`.  If `density = True`, call `ndimage.gaussian_filter` on the result of `np.histogram2d` with smoothing parameter `sigma`.  This will create a smoothed version of the raw data histograms.

```python
def bin_shots(df, bin_edges, density=False, sigma=1):
    
    """Given data frame of shots, compute a 2d matrix of binned counts is computed
    
    Args:
        df: data frame of shotchartdetail from nba.com. 
            At the minimum, variables named LOCX and LOCY are required.
        bin_edges: bin edge definition: edges in x and edges in y
    
    Returns:
        binned: counts
        xedges: bin edges in X direction
        yedges: bin edges in Y direction
    """
    import numpy as np
    from scipy import ndimage
    
    ## Call np.histogram2d
    binned, xedges, yedges = ...
    
    if density:

        # Recompute 'binned' using "gaussian_filter"
        binned = ...
        
        # Normalize the histogram to be a "density", e.g. mass across all bins sums to 1.
        binned /= np.sum(binned)
    
    return(binned, xedges, yedges)

```

```python
grader.check("q3a")
```

<!-- BEGIN QUESTION -->

### Question 3b: Visualize the binning on `curry_data`

Call `bin_shots` on `curry_data` to create a binned but unsmoothed matrix of shot counts (call this `curry_binned_unsmoothed`), a binned and smoothed matrix of counts with `sigma=1` (call this `curry_binned_smoothed1`) and one with `sigma=5` (call this `curry_binned_smoothed5`). Use the bin edges defined below:

```python
## bin edge definitions in inches
xedges = np.linspace(start=-300, stop=300, num=151)
yedges = np.linspace(start=-48, stop=372, num=106)
```

_Type your answer here, replacing this text._

```python
bin_edges = (xedges, yedges)

# curry_binned_unsmoothed, xe, ye = ... 
# curry_binned_smoothed1,  xe, ye = ... 
# curry_binned_smoothed5,  xe, ye = ... 
...
```

```python
def plot_shotchart(binned_counts, xedges, yedges, ax=None, use_log=False, cmap = 'Reds'):
    
    """Plots 2d heatmap from vectorized heatmap counts
    
    Args:
        hist_counts: vectorized output of numpy.histogram2d
        xedges, yedges: bin edges in arrays
        ax: figure axes [None]
        use_log: will convert count x to log(x+1) to increase visibility [False]
        cmap: Set the color map https://matplotlib.org/examples/color/colormaps_reference.html
    Returns:
        ax: axes with plot
    """
    
    import numpy as np
    import matplotlib.pyplot as plt
    
    ## number of x and y bins.  
    nx = xedges.size - 1
    ny = yedges.size - 1

    X, Y = np.meshgrid(xedges, yedges)
    
    if use_log:
        counts = np.log(binned_counts + 1)
        
    if ax is None:
        fig, ax = plt.subplots(1,1)
    
    ax.pcolormesh(X, Y, binned_counts.T, cmap=cmap) 
    ax.set_aspect('equal')
    
    draw_court(ax)
    
    return(ax)
```

```python
fig, ax = plt.subplots(1, 3, figsize=(20,60))

...
fig.show()
```

<!-- END QUESTION -->

### Vectorize Shot Images

- Here we proceed create a dictionary of smoothed patterns, each vectorized into a 1-d array (like Lab 6)
- In this case, the object `all_smooth` is a dictionary that consists of arrays of length `15750`. 
- Each entry in `all_smooth` represents the smoothed frequency of shots along the bins generated in the code above for a given player.

```python
## number of bins is one less than number of edges (remember homework 1)
nx = xedges.size - 1
ny = yedges.size - 1

## 2d histogram containers for binned counts and smoothed binned counts
all_counts = []
all_smooth = []
pids = []

## 2d histogram containers for binned counts and smoothed binned counts

## data matrix: players (row) by vectorized 2-d court locations (column)
for i, one in enumerate(allshots.groupby('PLAYER_ID')):
    
    ## what does this line do?
    pid, pdf = one
        
    num_shots = len(pdf.index)
    if(num_shots > 100): 
        
        tmp1, xedges, yedges = bin_shots(pdf, bin_edges=(xedges, yedges), density=True, sigma=2)
        tmp2, xedges, yedges = bin_shots(pdf, bin_edges=(xedges, yedges), density=False)
    
        ## vectorize and store into list
        all_smooth += [tmp1.reshape(-1)]
        all_counts += [tmp2.reshape(-1)]
        pids += [pid]
        
X = np.vstack(all_smooth).T
p, n = X.shape

print('Number of shot regions (p):', p)
print('Number of players (n):', n)
```

## Question 4: Non-negative Matrix Factorization (NMF) 

The non-negative matrix factorization is a dimension reduction technique that is often applied to image data.  It is similar to PCA except that is only applicable for strictly positive data.  We can apply the NMF to vectorized versions of the shot surface. This is useful because we can convert the observed matrix of shot surfaces into:
* Bases: Identifying modes of shooting style (number of modes is determined by `n_components` argument to `NMF` function below)
* Coefficients: How each players shooting style could be expressed as a (positive) linear combination of these bases

The NMF solves the following problem: given some matrix $X$ is $p\times n$ matrix, NMF computes the following factorization:
$$ \min_{W,H} \| X - WH \|_F\\
\text{ subject to } W\geq 0,\ H\geq 0, $$
where $W$ is ${p\times r}$ matrix and $H$ is ${r\times n}$ matrix.

In this homework, we have the following:
 
#### The data matrix $X$ 
$X$ is of dimension $n$={number of players} and $p$={number of total square bins on the court}. Each column corresponds to a player, with entries corresponding to a "flattened" or "vectorized" version of the 2d histograms plotted in part 4b.

#### Bases matrix: $W$

Columns $W_i$ contain the shot "bases". First, we will try it with $r=3$ bins in 5a, and then with $r=10$ bins in 5d.

#### Coefficient matrix: H

Each column of $H$ gives a coefficient for each of the bases vectors in $W$, and there are $n$ columns for each player.


The `sklearn` library is one of the main Python machine learning libraries.  It has a built in NMF function for us.  The function below runs this function and normalizes the basis surfaces to sum to 1.

```python
## Non-negative Matrix Factorization
def non_negative_marix_decomp(n_components, array_data):
    import sklearn.decomposition as skld
    model = skld.NMF(n_components=n_components, init='nndsvda', max_iter=500, random_state=0)
    W = model.fit_transform(array_data)
    
    # Normalize basis vectors to sum to 1
    Wsum = W.sum(axis=0)
    W = W/Wsum
    
    ## fix H correspondingly
    H = model.components_
    H = (H.T * Wsum).T

    nmf = (W, H)
    return(nmf)
```

### Question 4a: Computing NMF Factorization

Compute the NMF on all player's shot charts, X, assuming with `n_components` = 3 (i.e. each shot chart can be represented as a positive linear combination of 3 "basis" shot charts).  Fill in `plot_vectorized_shot_chart`.  This takes a the a vector of binned shot counts, converts it back to a matrix of the appropriate size and then calls `plot_shotchart` on the matrix.  The numpy function `reshape` will be useful here: [https://docs.scipy.org/doc/numpy/reference/generated/numpy.reshape.html](https://docs.scipy.org/doc/numpy/reference/generated/numpy.reshape.html)


```python
# W3, H3 = ...
```

```python
grader.check("q4a")
```

<!-- BEGIN QUESTION -->

### Question 4b: Visualizing Shot Types

Plot the first three basis images by calling `plot_vectorized_shot_chart` below on the columns of `W3`.

```python
def plot_vectorized_shotchart(vec_counts, xedges, yedges, ax=None, use_log=False, cmap = 'Reds'):
    
    """Plots 2d heatmap from vectorized heatmap counts
    
    Args:
        hist_counts: vectorized output of numpy.histogram2d
        xedges, yedges: bin edges in arrays
        ax: figure axes [None]
        use_log: will convert count x to log(x+1) to increase visibility [False]
        cmap: Set the color map https://matplotlib.org/examples/color/colormaps_reference.html
    Returns:
        ax: axes with plot
    """

    nx = ...
    ny = ...
    
    # use reshape to convert a vectorized counts back into a 2d histogram
    two_d_counts = ...
    
    return(plot_shotchart(two_d_counts, xedges, yedges, ax=ax, use_log=use_log, cmap=cmap))

fig, ax = plt.subplots(1, 3, figsize=(20,60))

## Write a for loop
for i in range(3):
    # Call plot_vectorized_shot_chart
    ax[i].set_title('Shot Basis %i' % (i+1))
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 4c: Reconstruction Error

Below we re-construct the shooting pattern for a single player. By "reconstructing" we mean use the approximation $$\hat{X} = WH$$ obtained via NMF.  Find $\hat X$ by multipling W and H.  In python the `@` symbol is used for matrix multiplication.


_Type your answer here, replacing this text._

```python
X3_hat = ...
```

```python
# Find the player_id of LaMarcus Aldridge
player_id = ...

## find index in X corresponding to that player
to_plot_idx = np.where(pids == player_id)[0][0]

fig, ax = plt.subplots(1, 3, figsize=(20,60))

## Call plot_vectorized_shot_chart
original_shotchart = ...
reconstructed_shotchart = ...
residual_chart = ...

# print(max(abs(X3_hat[:, to_plot_idx] - X[:, to_plot_idx])))
ax[0].set_title('Original Shooting Pattern')
ax[1].set_title('Reconstructed Shooting pattern (r=3)')
ax[2].set_title('Residual Shooting Pattern (r=3)')
None # prevents the title string from showing up as output
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 4d: Choice of Colormap

Why does it make sense to use a _sequential_ palette for the original and reconstructed shot charts and a _diverging_ palette for the residual?  _Hint:_ Read the introduction to colormaps [here](https://matplotlib.org/users/colormaps.html). 


**SOLUTION**

What areas of the court does this player to shoot more and where less relative to the reconstructed area.  If its helpful, you can refer to court locations by name using this legend [here](https://en.wikipedia.org/wiki/Basketball_court#/media/File:Basketball_terms.png]).

**SOLUTION**

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 4e: More Detailed Modeling

Re-run the analysis, this time for 10 basis vectors instead of 3.  Again plot the bases using `plot_vectorized_shotchart` on the columns of `W10`.

**Hint**: Study the following code
```
fig, ax = plt.subplots(2, 5, figsize=(20, 7))
ax = ax.flatten() # turn ax into a flat array
ax[0].set_title('hello')
ax[9].set_title('there')
fig.show()
```

```python
W10, H10 = ...

fig, ax = plt.subplots(2, 5, figsize=(20, 7))

## Write a for loop
for i in range(10):
    ...
    ax[i//5, i % 5].set_title('Shot Basis %i' % (i+1))

```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

If you did things correctly, you should be really impressed! We've identified potentially interesting patterns of shooting styles without actually specifying anything about the way basketball is played or where the relevant lines are on the court.  The resulting images are based only on the actual behavior of the players.  Even more impressive is that we're capturing similarity in regions that are far apart on the court.  One reason we can do this is that a basketball court is symmetric along the length of the court (i.e. symmetric about x=0).  However, people tend to be left or right hand dominant, which might affect their preferences.  Look carefuly at the shot basis plots above: is there any evidence of _asymmetry_ in player shooting behavior?  Refer to specific basis images in your answer.    


**SOLUTION**

<!-- END QUESTION -->

Repeat part 5b, and again plot original, reconstructed and residual shot chats for LaMarcus Aldridge.


<!-- BEGIN QUESTION -->



```python
X10_hat = ...

fig, ax = plt.subplots(1, 3, figsize=(20,60))

# I took the first player appearing in first column 
# (you probably want to do more interesting players)
original_shotchart = ...
reconstructed_shotchart = ...
residual_chart = ...

ax[0].set_title('Original Shooting Pattern')
ax[1].set_title('Reconstructed Shooting pattern (r=10)')
ax[2].set_title('Residual Shooting Pattern (r=10)');
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 4f: Comparing Players 

With `H10` matrix, it is possible to compare any pair of players. For all players pairwise, $i$ and $j$, compare using euclidean distance between their coefficients:

$$ \text{player-distance}(i,j) = \| H_i - H_j \|_2 = \left(\sum_{k=1}^{10} (H_{ki} - H_{kj})^2 \right)^{1/2} $$

Create a heatmap for comparing pair-wise player distance matrix. Find the two pairs of players with the smallest distances. Also, find two pairs of players with largest distances.

Hint: you can construct the distance matrix manually or use `scipy.spatial.distance_matrix`.

**SOLUTION**

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 4g: Residuals

The residual betwene `Xhat` and `X` gives a sense of how well a player is decribed by NMF computed matrices `W` and `H`. Calculate RMSE for each player, and plot the histogram. Comment on this distribution and players with smallest and largest RMSEs (use 10 components).


**SOLUTION**

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 4h: Proposing improvements

One of the main purposes of exploratory data analysis is to generate new ideas, directions, and hypothesis for future analyses and experiments. Take two players of your choice and compare their shooting patterns with various visualizations.

State any insights and defend your conclusions with visual and/or numerical comparisons.

**SOLUTION**

<!-- END QUESTION -->

_Cell Intentionally Blank_

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Download the zip file and submit to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1a (points: 3)
```
>>> ## TEST ##
>>> print(type(allteams))
<class 'pandas.core.frame.DataFrame'>
```
```
>>> ## TEST ##
>>> allteams.columns 
Index(['LEAGUE_ID', 'TEAM_ID', 'MIN_YEAR', 'MAX_YEAR', 'ABBREVIATION'], dtype='object')
```
```
>>> ## TEST ##
>>> allteams.shape 
(30, 5)
```


Tests for q1b (points: 3)
```
>>> ## TEST ##
>>> print(type(allplayers))
<class 'pandas.core.frame.DataFrame'>
```
```
>>> ## TEST ##
>>> allplayers.columns 
Index(['DISPLAY_LAST_COMMA_FIRST', 'DISPLAY_FIRST_LAST', 'ROSTERSTATUS',
       'FROM_YEAR', 'TO_YEAR', 'PLAYERCODE', 'TEAM_ID', 'TEAM_CITY',
       'TEAM_NAME', 'TEAM_ABBREVIATION', 'TEAM_CODE', 'GAMES_PLAYED_FLAG',
       'OTHERLEAGUE_EXPERIENCE_CH'],
      dtype='object')
```
```
>>> ## TEST ##
>>> allplayers.shape 
(4540, 13)
```
```
>>> ## TEST ##
>>> allplayers.index.name 
'PERSON_ID'
```
```
>>> ## TEST ##
>>> allplayers.tail(3).DISPLAY_LAST_COMMA_FIRST
PERSON_ID
1629967    Flatten, Skyler
1630001       Morgan, Matt
1630003        Scott, MIke
Name: DISPLAY_LAST_COMMA_FIRST, dtype: object
```


Tests for q1c (points: 4)
```
>>> ## TEST ##
>>> len(allshots_files)
502
```
```
>>> ## TEST ##
>>> len(allshots_list)
502
```
```
>>> ## TEST ##
>>> all([type(df)==pd.core.frame.DataFrame for df in allshots_list])
True
```
```
>>> ## TEST ##
>>> print([len(df) for df in allshots_list])
[36, 652, 720, 135, 117, 578, 104, 311, 381, 1138, 157, 120, 362, 740, 446, 157, 888, 558, 13, 256, 1517, 1314, 365, 749, 741, 842, 298, 1255, 780, 495, 440, 220, 628, 347, 117, 487, 666, 400, 217, 104, 0, 439, 350, 295, 1026, 191, 479, 464, 13, 343, 844, 179, 260, 112, 26, 0, 87, 48, 960, 197, 700, 737, 180, 22, 506, 95, 1360, 729, 74, 49, 483, 268, 0, 1173, 290, 594, 0, 792, 616, 748, 271, 424, 63, 181, 483, 945, 0, 98, 245, 218, 331, 352, 127, 81, 217, 314, 77, 379, 528, 600, 215, 418, 791, 252, 324, 426, 173, 262, 275, 18, 16, 155, 456, 841, 1102, 1036, 586, 611, 650, 282, 795, 1530, 520, 400, 831, 485, 404, 348, 568, 183, 486, 431, 251, 41, 14, 192, 16, 212, 1087, 18, 158, 543, 447, 53, 464, 73, 257, 21, 64, 274, 169, 162, 49, 122, 264, 106, 502, 702, 32, 698, 312, 12, 283, 8, 141, 76, 10, 266, 386, 396, 362, 0, 8, 20, 16, 0, 27, 39, 178, 3, 706, 243, 142, 563, 511, 314, 234, 565, 185, 371, 165, 124, 3, 53, 21, 716, 216, 262, 657, 209, 589, 16, 0, 914, 4, 256, 58, 87, 140, 508, 0, 110, 291, 1206, 557, 63, 34, 244, 69, 50, 518, 19, 171, 46, 279, 1256, 870, 1186, 73, 47, 39, 4, 2, 51, 22, 290, 36, 449, 33, 33, 60, 151, 5, 0, 13, 110, 9, 39, 82, 0, 46, 27, 0, 51, 110, 3, 126, 23, 0, 18, 8, 33, 26, 78, 10, 0, 21, 47, 468, 376, 1319, 746, 1027, 149, 432, 739, 523, 665, 385, 1383, 723, 1120, 687, 123, 213, 841, 690, 239, 273, 870, 259, 753, 1473, 283, 884, 938, 677, 786, 230, 535, 641, 69, 129, 153, 877, 562, 376, 446, 433, 1341, 1909, 660, 730, 1340, 1313, 379, 1159, 416, 659, 537, 575, 457, 158, 355, 39, 477, 630, 663, 493, 698, 552, 443, 619, 128, 371, 72, 189, 593, 635, 1614, 302, 219, 971, 608, 58, 680, 593, 432, 0, 541, 489, 197, 1241, 683, 380, 538, 156, 260, 1684, 29, 1402, 474, 484, 844, 1129, 1354, 446, 1254, 6, 265, 542, 1047, 548, 904, 1051, 403, 182, 113, 453, 532, 99, 105, 1026, 332, 1609, 478, 1533, 1027, 1052, 1027, 572, 242, 979, 67, 388, 30, 1069, 105, 48, 530, 797, 422, 1148, 460, 420, 484, 136, 691, 777, 476, 223, 276, 648, 373, 0, 1243, 345, 1104, 526, 377, 564, 756, 442, 123, 361, 643, 594, 362, 712, 809, 993, 361, 548, 588, 1247, 294, 242, 178, 161, 383, 465, 1, 275, 461, 350, 438, 35, 435, 130, 1135, 412, 1180, 637, 828, 338, 671, 151, 176, 823, 748, 422, 1046, 611, 567, 276, 486, 440, 1089, 1209, 749, 1199, 241, 54, 7, 696, 732, 863, 826, 225, 1206, 0, 526, 261, 587, 802, 419, 99, 94, 463, 116, 1095, 960, 193, 483, 33, 69, 210, 347, 118, 302, 79, 736]
```
```
>>> ## TEST ##
>>> type(allshots) is pd.DataFrame 
True
```
```
>>> ## TEST ##
>>> allshots.shape 
(217317, 23)
```
```
>>> ## TEST ##
>>> print(allshots.index)
Index([   1713,    1713,    1713,    1713,    1713,    1713,    1713,    1713,
          1713,    1713,
       ...
       1629541, 1629541, 1629541, 1629541, 1629541, 1629541, 1629541, 1629541,
       1629541, 1629541],
      dtype='object', name='PLAYER_ID', length=217317)
```


Tests for q1d (points: 6)
```
>>> ## TEST ##
>>> type(curry_id) == str
True
```
```
>>> ## TEST ##
>>> print(type(curry_data))
<class 'pandas.core.frame.DataFrame'>
```
```
>>> ## TEST ##
>>> curry_data.shape
(1340, 23)
```
```
>>> ## TEST ##
>>> curry_data.SHOT_MADE_FLAG.dtype
dtype('bool')
```


Tests for q2a (points: 5)
```
>>> ## TEST ##
>>> ax2a.axes.get_ylim()
(-100.0, 500.0)
```
```
>>> ## TEST ##
>>> ax2a.axes.get_xlim()
(-300.0, 300.0)
```


Tests for q3a (points: 2)
```
>>> ## TEST ##
>>> 
>>> ## Smoothed surface should sum to 1
>>> xedges = np.linspace(start=-300, stop=300, num=151)
>>> yedges = np.linspace(start=-48, stop=372, num=106)
>>> bin_shots(pd.DataFrame({'LOC_X' : curry_data.LOC_X , 'LOC_Y' : curry_data.LOC_Y}), (xedges, yedges), density=True)[0].sum()
1.0
```
```
>>> ## TEST ##
>>> 
>>> ## Total number of shots in the binned results is right
>>> xedges = np.linspace(start=-300, stop=300, num=151)
>>> yedges = np.linspace(start=-48, stop=372, num=106)
>>> bin_shots(pd.DataFrame({'LOC_X' : curry_data.LOC_X , 'LOC_Y' : curry_data.LOC_Y}), (xedges, yedges))[0].sum()
1328.0
```
```
>>> ## TEST ##
>>> 
>>> ## Number of bined shots in smaller xrange works
>>> xedges2 = np.linspace(start=-100, stop=100, num=151)
>>> yedges = np.linspace(start=-48, stop=372, num=106)
>>> bin_shots(pd.DataFrame({'LOC_X' : curry_data.LOC_X , 'LOC_Y' : curry_data.LOC_Y}), (xedges2, yedges))[0].sum()
589.0
```
```
>>> ## TEST ##
>>> 
>>> ## Number of bined shots in smaller yrange works
>>> xedges = np.linspace(start=-300, stop=300, num=151)
>>> yedges2 = np.linspace(start=0, stop=100, num=151)
>>> bin_shots(pd.DataFrame({'LOC_X' : curry_data.LOC_X , 'LOC_Y' : curry_data.LOC_Y}), (xedges, yedges2))[0].sum()
479.0
```


Tests for q4a (points: 5)
```
>>> ## TEST ##
>>> type(H3) is np.ndarray
True
```
```
>>> ## TEST ##
>>> W3.shape
(15750, 3)
```
```
>>> ## TEST ##
>>> H3.shape
(3, 388)
```
